# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">PHYS 231 Lab #4 - Day 1<span style="color:red"> $\to$ (8 possible marks)</span></span></div>

### Due Thursday, Oct. 19 at 17:00 $-$ upload your .zip submission to the PHYS 231 Canvas Gradebook.

<font size ="4"><span style="color:purple">We recommend that you complete this notebook using either the [Google Chrome](https://www.google.com/intl/en_ca/chrome/) or [Mozilla Firefox](https://www.mozilla.org/en-CA/firefox/new/) browser.  Chrome and Firefox are both available for Windows, Mac, and Linux operating systems.  The PHYS 121 Jupyter notebooks should work in other browsers, but we have specifically verfied that they work well in both Chrome and Firefox.</span></font>

# $LRC$ Circuit Transient Response

***
## Lab Learning Outcomes:
By the end of the PHYS 231 Lab #4, students will be able to:
* <b><span style="color:rgba(0, 153, 76, 1);">Design, build, and test practical $LRC$ circuits.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Determine the time constant from the transient response of an $LRC$ circuit.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Perform a weighted linear fit to a set of data and use the extracted best-fit parameters to determine a quantity of interest.</span></b>

***
## Autograding:
The PHYS 231 Labs will make use of some autograding.  To make the autograding work, the two cells below needs to be executed.  The first one installs the required packages and the second imports required packages/modules.  If 'PHYS231.Installer()' reports that some functions have been installed, the user should restart the "kernel" and then exectute the PHYS231.Installer() cell a second time.  The second time the installer function is run, it should report that **"All packages already installed. Please proceed"**.

If necessary, the kernel can be restarted by selecting **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/restartKernel.gif">GIF</a> showing how to restart the kernel.

The 'PHYS231.Installer()' command requires the file 'PHYS231.py', which you should see included in the list of files along the left-hand side of the screen.

```python
# Hide unnecessary warnings.
import warnings
warnings.filterwarnings('ignore')

# Import PHYS231.py and then run the installer function.
import PHYS231
PHYS231.Installer()
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 231 - Lab 4a.ipynb")
```

***
## Import Modules:
Execute the cell below to import a number of useful pre-built Python modules that will be used in the PHYS 121 pre-labs and labs.

```python
import numpy as np
import matplotlib.pyplot as plt
from matplotlib.pyplot import cm # used to generate a sequence of colours for plotting
import pandas as pd
import scipy.stats as sp
from scipy.optimize import curve_fit
from IPython.display import HTML as html_print
from IPython.display import display, Markdown, Latex
import uncertainties
import math
import sys
```

```python
PHYS231.graderCheck(['PHYS231' in sys.modules, 'PHYS231' in sys.modules, 'numpy' in sys.modules], ['PHYS231', 'otter', 'numpy'], grader.check('q0.1'))
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Name & Student Number</span></div>  

To ensure that both you and your partner receive credit for you lab work, both your name and student number and your partner's name and student number need to appear in the notebooks used during the labs.  In the cells below this text, enter your name and student number and your partner's name and student number as strings (i.e. place your names and student numbers between single or double quotes.

Examples of valid entries are:
```python
last_name_1 = 'Doe'
first_name_1 = 'Jane'
student_number_1 = '12345678'
```
and
```python
last_name_1 = "Doe"
first_name_1 = "John"
student_number_1 = "12345678"
```

<!-- BEGIN QUESTION -->

***
### **<span style="color:blue">Partner #1</span>** **<span style="color:red">(0.5 marks)</span>** 

```python
# Enter LAST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_1 = ...
```

```python
# Enter FIRST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_1 = ...
```

```python
# Enter STUDENT NUMBER of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_1 = ...
```

***
### **<span style="color:blue">Partner #2</span>** **<span style="color:red">(0.5 marks)</span>** 
If you are working on your own, leave all of last_name_2, first_name_2, and student_number_2 as ...  Please still exectute these cells and then run the 'grader.check("name_and_student_number_2")' cell.

```python
# Enter LAST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_2 = ...
```

```python
# Enter FIRST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_2 = ...
```

```python
# Enter STUDENT NUMBER of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_2 = ...
```

```python
PHYS231.graderCheck([last_name_1, first_name_1, student_number_1, last_name_2, first_name_2, student_number_2], ['last_name_1', 'first_name_1', 'student_number_1', 'last_name_2', 'first_name_2', 'student_number_2'], grader.check('name_and_student_number'))
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Oscilloscope Measurement of the Time Constant</span></div>

## NumPy Arrays

Data that you collect during the lab can be entered directly into a list using the following notation:

```python
x = [x0, x1, x2, x3, ..., xN]
y = [y0, y1, y2, y3, ..., yN]
```
These lists will contain $N+1$ elements. It's usually more convenient to enter data into a NumPy array rather than a list because you can do mathematical operations on arrays that you can't do on lists.  For example:

- squaring an array creates an new array with each of the elements from the original array squared
- multiplying two arrays of the same length creates a new array in which the each element is the product of the corresponding elements from the original arrays
- ...

NumPy arrays are entered using the following syntax:

```python
x = np.array([x0, x1, x2, x3, ..., xN])
y = np.array([y0, y1, y2, y3, ..., yN])
```             

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.1</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with your data.  You have **lots** of time in today's lab, so focus on collecting as much high-quality data as possible.

*** Please do not change anything to the left of the equals sign. ***

```python
time = np.array([...])
Vp = np.array([...])
errVp = np.array([...])
```

```python
PHYS231.graderCheck([time, Vp, errVp], ['time', 'Vp', 'errVp'], grader.check('q1.1'))
```

## Pandas DateFrames

The Pandas package can be used to display your array data in a nicely-formated table.  The Pandas package creates "DataFrames" which are two-dimensional data structures, like a two-dimensional array, or a table with rows and columns.  The cell below creates and then displays a DataFrame using the arrays that you created in the cell above.

```python
# Create the data structure.
data = {
    "time (s)": time,
    "VR (V)": Vp,
    "errVR (V)": errVp
        }

# An option to format the numbers in the DataFrame.
pd.set_option('display.float_format', lambda x: '%0.3g' % x)

#Display the data.
pd.DataFrame(data)
```

### Calculate the Natural Logarithm of the Peak Voltage

The transient response of voltage a across are resistor in an highly-underdamped $LRC$ circuit is given by:

\begin{align}
V_\mathrm{R}(t)\approx\omega_1 V_0 RCe^{-\left(\gamma/2\right)t}\sin\left(\omega_1 t\right)
\end{align}

where $V_0$ is the size of the voltage step, $\gamma = R/L$, and $\omega_1=\sqrt{\omega_0^2 - \left(\gamma/2\right)^2}$.  Note that, for the highly-underdamped circuit that we're working with $\omega\gg\gamma/2$ such that $\omega_1\approx\omega_0=1/\sqrt{LC}$. In this case, the peak voltages as a function of time are governed by:

\begin{align}
V_\mathrm{p}(t)\approx\omega_0 V_0 RCe^{-\left(\gamma/2\right)t}
\end{align}

This expression can be linearized by taking the natural logarithm of both sides such that:

\begin{align}
\ln\left(V_\mathrm{p}\right)\approx\left(-\frac{\gamma}{2}\right) t + \ln\left(\omega_0 V_0 RC\right)
\end{align}

This is the equation of a straight line with slope $-\gamma/2$ and $y$-intercept $\ln\left(\omega_0 V_0 RC\right)$.

In the cell below, you will calculate the natural logarithm of the peak voltage.  Recall that in Python, the natural logarithm of $x$ is called using
```python
np.log(x)
```

***
**<span style="color:blue">Question 1.2</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with an expression that calculates the $ln\left(V_\mathrm{p}\right)$.

*** Please do not change anything to the left of the equals sign. ***

```python
lnVp = ...
```

```python
PHYS231.graderCheck([lnVp], ['lnVp'], grader.check('q1.2'))
```

### Calculate the uncertainty in $\ln\left(V_\mathrm{p}\right)$

Recall propagation of errors.  If $y = f\left(x1, x2, ..., x_N\right)$ and $x_i\pm\Delta x_i$ for $i= 1...N$ are known, then:

\begin{align}
\Delta y = \sqrt{\left(\frac{\partial f}{\partial x_1}\Delta x_1\right)^2 + \left(\frac{\partial f}{\partial x_2}\Delta x_2\right)^2 + ... + \left(\frac{\partial f}{\partial x_N}\Delta x_N\right)^2}
\end{align}

If $y=\ln\left(V_\mathrm{p}\right)$, then $y$ is a function of only a single variable and:

\begin{align}
\Delta y = \left\vert \frac{\partial y}{\partial V_\mathrm{p}}\Delta V_\mathrm{p}\right\vert
\end{align}

In the cell below, calculate the uncertainty in $\ln V_\mathrm{p}$.

***
**<span style="color:blue">Question 1.3</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with an expression that calculates the uncertainty in $ln\left(V_\mathrm{p}\right)$.

*** Please do not change anything to the left of the equals sign. ***

```python
errlnVp = ...
```

```python
PHYS231.graderCheck([errlnVp], ['errlnVp'], grader.check('q1.3'))
```

### Plot the data

The $\ln\left(V_\mathrm{p}\right)$ data can now be plotted as a function of time.  Execute the cell below to created a scatter plot with error bars.  For additional information about the 'Scatter' function, see **Appendix A** near the end of this notebook. 

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.4</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with the appropriate data to generate a plot of $ln\left(V_\mathrm{p}\right)$ versus $t$.  Be sure to include errror bars.

For additional information about the 'Scatter()' function, see **Appendix A** near the end of this notebook. 

*** Please do not change anything to the left of the equals sign. ***

```python
PHYS231.Scatter(xData = ..., 
                yData = ..., 
                yErrors = ...,
                xlabel = 'time', 
                ylabel = 'ln Vp', 
                xUnits = 's', 
                yUnits = '');
```

<!-- END QUESTION -->

## Linear Fits

The final thing that we'll do in Part 2 is fit a straight line through the data corrected (i.e. cut) data.  We'll use the 'LinearFit' contained in PHYS231.py to complete this task.  The function call to perform a weighted fit is:
```python
m, b, dm, db, fig = PHYS231.LinearFit(t, y, erry)
```
A weighted fit with a formatted plot can be called by using some of the other available arguments:
```python
m, b, dm, db, fig = PHYS231.LinearFit(xData = t, yData = y, yErrors = erry, xlabel = 'time', ylabel = 'position', xUnits = 's', yUnits = 'm')
```
In the cell below to preform a weighted linear fit to your data.

For additional information about the 'LinearFit' function, see **Appendix B** near the end of this notebook. 

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.5</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with the data needed to a weighted linear fit to $\ln\left(V_\mathrm{p}\right)$ versus $t$.

*** Please do not change anything to the left of the equals sign. ***

```python
m, b, dm, db, fig = PHYS231.LinearFit(
                    xData = ...,
                    yData= ...,
                    yErrors = ...,
                    xlabel = "time",
                    ylabel = "ln Vp",
                    xUnits = "s",
                    yUnits = "",
)
```

<!-- END QUESTION -->

## Determining the time constant: $\gamma\pm\Delta\gamma$

Use the slope from the linear fit and its uncertainty to determine an experimental value for $\gamma=R/L$ and its uncertainty.

***
**<span style="color:blue">Question 1.6</span>** **<span style="color:red">(2 marks)</span>**  

In the cell below, replace the ... with expressions that calculate $\gamma$ and its uncertainty.

*** Please do not change anything to the left of the equals sign. ***

```python
# Enter an expression to calculate gamma. 
gamma = ...

# Enter an expression to calculate the uncertainty in gamma. 
err_gamma = ...
```

```python
PHYS231.graderCheck([gamma, err_gamma], ['gamma', 'err_gamma'], grader.check('q1.6'))
```

```python
# The 'uncertainties' package can be used round the numerical value of gamma based on the value of the uncertainty.
# gamma and errgamma are passed to 'ufloat'
print('gamma =', uncertainties.ufloat(gamma, err_gamma), '1/s')
```

In your lab notebook, you should compare this experimentally-measured value of $\gamma$ to the value that you expected based on the values of the components used in your circuit.  Recall that $\gamma=R/L$.  Keep in mind that your expected value of $\gamma$ will also have an uncertainty based on the uncertainties in the values of $R\pm\Delta R$ and $L\pm\Delta L$.

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Feedback and Submission</span></div>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.1:</span>**  

We welcome your feedback on the PHYS 231 labs!  Please feel free to include any comments you have about this lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 231 labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the lab, ...

This question is optional and will have no impact on your lab grade.

***
**<span style="color:blue">Answer 2.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell to enter your comments.  This text can be deleted.

<!-- END QUESTION -->

***
Once you've completed this notebook:
- Save your work.
- Run 'grader.check_all()' to confirm that you've completed all required tasks.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file using the link that appears **below** the ```grader.check_all(...)``` cell.
- $\Leftarrow$ Please do **NOT** download the zip file from the frame along the left-hand side of the screen. 
- Upload the .zip file to the PHYS 121 Canvas gradebook.
- **Do NOT change the name of the .zip file.**
- **Do NOT modify the contents of the .zip file.**

Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
display(grader.check_all())
PHYS231.saveMessage()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
PHYS231.waitMessage()
grader.export(files=["PHYS121_DataLogger.txt"])
PHYS231.proceedMessage()
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  

```python
# Here's an empty code cell that you can use.
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix A &ndash; The Scatter Function...</div>

The function for generating scatter plots is called as follows:
```python
PHYS231.Scatter(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False, show = True)
```
The 'xData' and 'yData' inputs are required, all other arguments are optional with default values set.  The function returns the a single output (the formatted plot):
```python
fig
```

The function will do a simple scatter plot if no 'yError' are included.  It will included error bars if 'yErrors' are passed to the function.  The 'fill' and 'show' arguments should generally be false.  

If the 'xData' list is empty, as in:
```python
xData = []
```
then the x-axis will be the trial number.  For example, is 'xData' is empty and 'yData' has 10 elements, the x-axis will span 1 to 10.

### Scatter Example Implmentation
The code block below shows an implementation of 'Scatter'.
```python
import PHYS231
theta = [10, 20, 30] # degrees
T = [2.02, 1.95, 2.13] # s
errT = [0.02, 0.02, 0.03] # s
fig = PHYS231.Scatter(theta, T, errT, 'initial angle' , 'period', 'degrees', 's')
```

If you're interested in generating your own plots with customized formatting, see the following Python-based plotting tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Basic%20Plots.html.
***

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix B &ndash; The LinearFit Function...</div>

The linear-fit function is called as follows:
```python
PHYS231.LinearFit(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False)
```
The xData and yData lists are required.  The others are optional with default values set.  The function returns the following outputs:
```python
Slope, Yintercept, errSlope, errYintercept, fig
```
The first four outputs are numerical values and 'fig' is the formatted plot. 

If $y$-uncertainities are provided, the function will perform a weighted fit.  The $y$-uncertainties list must be the same length as $x$- and $y$-data lists.  If $y$-uncertainties are not provided, the fit will be unweighted.  The other optional arguments include $x$- and $y$-axis names and units for the $x$- and $y$-datasets.  These must be entered as strings (enclosed in quotations) and they are used for formatting the outputs of the function.  If 'fill' is set to True, shading will be added around the best-fit line representing confindence intervals. 

### Unweighted Fit Example Implmentation
The code block below shows the most basic use of 'LinearFit' for an unweighted fit without any of the optional arguments.
```python
import PHYS231
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
m, b, dm, db, fig = PHYS231.LinearFit(V, I)
```

### Weighted Fit Example Implmentation
This second blcok of code shows how to use 'LinearFit' to do a weighted fit.  It also makes use of all of the other optional arguments.  
```python
import PHYS231
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
errI = [0.005, 0.012, 0.020, 0.025]
m, b, dm, db, fig = PHYS231.LinearFit(V, I, errI, 'voltage' , 'current', 'V', 'A', True)
```

The 'm, b, dm, db, fig = ...' syntax is used to store the function outputs in varibles.  If you're interested in details of the fitting process, see the following Python-based tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Weighted%20Linear%20Fit.html
***

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/ubc-logo-full.jpg" width="500"/>

Last update: February 12, 2025

Tests for name_and_student_number (points: None)
```
>>> isinstance(last_name_1, str)
True
```
failure_message: last_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_1, str)
True
```
failure_message: first_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_1, str)
True
```
failure_message: student_number_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = len(student_number_1) == 8
>>> test
True
```
failure_message: student_number_1 must be a string of exactly 8 characters.
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = student_number_1.isdigit()
>>> test
True
```
failure_message: student_number_1 must contain only interger digits (0, 1, 2, 3, ..., 9)
```
>>> isinstance(last_name_2, str) or last_name_2 == ...
True
```
failure_message: last_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_2, str) or first_name_2 == ...
True
```
failure_message: first_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_2, str) or student_number_2 == ...
True
```
failure_message: student_number_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = len(student_number_2) == 8
>>> test
True
```
failure_message: student_number_2 must be a string of exactly 8 characters or be left as ...
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = student_number_2.isdigit()
>>> test
True
```
failure_message: student_number_2 must contain only interger digits (0, 1, 2, 3, ..., 9) or be left as ...


Tests for q0.1 (points: None)
```
>>> 'PHYS231' in sys.modules and 'otter' in sys.modules and ('numpy' in sys.modules)
True
```
failure_message: Make sure you execute the cells containing "import PHYS231" and "import otter" as well as the cell beneath the heading "Import Modules".


Tests for q1.1 (points: None)
```
>>> len(time) == len(Vp)
True
```
failure_message: Your time and Vp lists are different lengths.
```
>>> len(time) == len(errVp)
True
```
failure_message: Your time and errVp lists are different lengths.
```
>>> len(Vp) == len(errVp)
True
```
failure_message: Your Vp and errVp lists are different lengths.


Tests for q1.2 (points: None)
```
>>> all(np.isclose(lnVp, np.log(Vp), rtol=0.01))
True
```
failure_message: Double check your calculation of ln(Vp).  Remember that, in Python, the syntax for the natural logarithm of x is np.log(x).


Tests for q1.3 (points: None)
```
>>> all(np.isclose(errlnVp, errVp / Vp, rtol=0.01))
True
```
failure_message: Double check your calculation of the uncertainty in ln(Vp).  You should apply propagation of errors.


Tests for q1.6 (points: None)
```
>>> math.isclose(gamma, -2 * m, abs_tol=0, rel_tol=0.01)
True
```
failure_message: Double check your calculation of gamma.
```
>>> math.isclose(err_gamma, 2 * dm, abs_tol=0, rel_tol=0.01)
True
```
failure_message: Double check your calculation of the uncertainty in gamma.
